<a href="https://colab.research.google.com/github/AdrianoDee/course_ml4hep/blob/main/notebooks/rnns/RNNs.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Machine translation using RNN

![](https://pytorch.org/tutorials/_images/seq2seq.png)

```python
from __future__ import unicode_literals, print_function, division
from io import open
import unicodedata



import string
import re
import random

import torch
import torch.nn as nn
from torch import optim
import torch.nn.functional as F

device = torch.device("cuda" if torch.cuda.is_available() else "cpu")
```

<!-- grade_id: b7efae, points: 0, grade: false, solution: true -->
```python
SOS_token = 0
EOS_token = 1


class Lang:
    def __init__(self, name):
        self.name = name
        self.word2index = {} #
        self.word2count = {} # words counter
        self.index2word = {0: "SOS", 1: "EOS"}
        self.n_words = 2  # Count SOS and EOS


    def addSentence(self, sentence):
        # Task: Add words splitting the sentence by space

    def addWord(self, word):
        # Task: Update self.word2index, self.word2count, self.n_words
```

<!-- grade_id: 2fea4b, points: 1, grade: true, solution: false -->
```python
test_lang = Lang('test')
test_lang.addSentence('Hello, World!')
assert test_lang.n_words == 4
assert test_lang.word2index == {'Hello,': 2, 'World!': 3}
assert test_lang.word2count == {'Hello,': 1, 'World!': 1}
```

```python
# Turn a Unicode string to plain ASCII, thanks to
# https://stackoverflow.com/a/518232/2809427
def unicodeToAscii(s):
    return ''.join(
        c for c in unicodedata.normalize('NFD', s)
        if unicodedata.category(c) != 'Mn'
    )

# Lowercase, trim, and remove non-letter characters

def normalizeString(s):
    s = unicodeToAscii(s.lower().strip())
    s = re.sub(r"([.!?])", r" \1", s)
    s = re.sub(r"[^a-zA-Z.!?]+", r" ", s)
    return s
```

```python
def readLangs(lang1, lang2, reverse=False):
    print("Reading lines...")

    # Read the file and split into lines
    lines = open('data/%s-%s.txt' % (lang1, lang2), encoding='utf-8').\
        read().strip().split('\n')

    # Split every line into pairs and normalize
    pairs = [[normalizeString(s) for s in l.split('\t')] for l in lines]

    # Reverse pairs, make Lang instances
    if reverse:
        pairs = [list(reversed(p)) for p in pairs]
        input_lang = Lang(lang2)
        output_lang = Lang(lang1)
    else:
        input_lang = Lang(lang1)
        output_lang = Lang(lang2)

    return input_lang, output_lang, pairs
```

```python
!wget https://download.pytorch.org/tutorial/data.zip
!unzip data.zip
```

Output:
```
--2023-06-15 07:38:55--  https://download.pytorch.org/tutorial/data.zip
Resolving download.pytorch.org (download.pytorch.org)... 13.224.249.33, 13.224.249.76, 13.224.249.92, ...
Connecting to download.pytorch.org (download.pytorch.org)|13.224.249.33|:443... connected.
HTTP request sent, awaiting response... 200 OK
Length: 2882130 (2.7M) [application/zip]
Saving to: ‘data.zip’

data.zip            100%[===================>]   2.75M  --.-KB/s    in 0.1s    

2023-06-15 07:38:55 (27.7 MB/s) - ‘data.zip’ saved [2882130/2882130]

Archive:  data.zip
   creating: data/
  inflating: data/eng-fra.txt        
   creating: data/names/
  inflating: data/names/Arabic.txt   
  inflating: data/names/Chinese.txt  
  inflating: data/names/Czech.txt    
  inflating: data/names/Dutch.txt    
  inflating: data/names/English.txt  
  inflating: data/names/French.txt   
  inflating: data/names/German.txt   
  inflating: data/names/Greek.txt    
  inflating: data/names/Irish.txt    
  inflating: data/names/It
```

```python
MAX_LENGTH = 10

eng_prefixes = (
    "i am ", "i m ",
    "he is", "he s ",
    "she is", "she s ",
    "you are", "you re ",
    "we are", "we re ",
    "they are", "they re "
)


def filterPair(p):
    return len(p[0].split(' ')) < MAX_LENGTH and \
        len(p[1].split(' ')) < MAX_LENGTH and \
        p[1].startswith(eng_prefixes)


def filterPairs(pairs):
    return [pair for pair in pairs if filterPair(pair)]
```

```python
def prepareData(lang1, lang2, reverse=False):
    input_lang, output_lang, pairs = readLangs(lang1, lang2, reverse)
    print("Read %s sentence pairs" % len(pairs))
    pairs = filterPairs(pairs)
    print("Trimmed to %s sentence pairs" % len(pairs))
    print("Counting words...")
    for pair in pairs:
        input_lang.addSentence(pair[0])
        output_lang.addSentence(pair[1])
    print("Counted words:")
    print(input_lang.name, input_lang.n_words)
    print(output_lang.name, output_lang.n_words)
    return input_lang, output_lang, pairs

# Uncomment if data does not exists
# !wget https://download.pytorch.org/tutorial/data.zip
# !unzip data.zip > /dev/null

input_lang, output_lang, pairs = prepareData('eng', 'fra', True)
print(random.choice(pairs))
```

Output:
```
Reading lines...
Read 135842 sentence pairs
Trimmed to 10599 sentence pairs
Counting words...
Counted words:
fra 4345
eng 2803
['tu es veinard .', 'you re in luck .']
```

<!-- grade_id: d3c5ff, points: 0, grade: false, solution: true -->
```python
def indexesFromSentence(lang, sentence):
    # split sentence by space and convert words to indices


def tensorFromSentence(lang, sentence):
    indexes = indexesFromSentence(lang, sentence)
    indexes.append(EOS_token)
    return torch.tensor(indexes, dtype=torch.long, device=device).view(-1, 1)


def tensorsFromPair(pair):
    input_tensor = tensorFromSentence(input_lang, pair[0])
    target_tensor = tensorFromSentence(output_lang, pair[1])
    return (input_tensor, target_tensor)
```

<!-- grade_id: 19371d, points: 1, grade: true, solution: false -->
```python
assert indexesFromSentence(output_lang, 'master driver') == [975, 977]
```

<!-- grade_id: dbc772, points: 0, grade: false, solution: true -->
```python
class EncoderRNN(nn.Module):
    def __init__(self, input_size, hidden_size):
        super(EncoderRNN, self).__init__()
        self.hidden_size = hidden_size

        self.embedding = nn.Embedding(input_size, hidden_size) # Special learnable layer which converts tensor of word indices to tensor of hidden size
        # Task:  Define your RNN here
        self.gru = nn.GRU(hidden_size, hidden_size)

    def forward(self, input, hidden):
        # 1: input is single sentence of word indexes
        # 2: PyTorch RNNs takes [seq_len, batch_size, hidden_size] tensor as input
        # Task: define forward operation here

        return output, hidden

    def initHidden(self):

        return torch.zeros(1, 1, self.hidden_size, device=device)
```

<!-- grade_id: cf9459, points: 0, grade: false, solution: true -->
```python
class DecoderRNN(nn.Module):
    def __init__(self, hidden_size, output_size):
        super(DecoderRNN, self).__init__()
        self.hidden_size = hidden_size

    def forward(self, input, hidden):
        return output, hidden

    def initHidden(self):
        return torch.zeros(1, 1, self.hidden_size, device=device)
```

<!-- grade_id: 63a789, points: 0, grade: false, solution: true -->
```python
def train(input_tensor, target_tensor, encoder, decoder, encoder_optimizer, decoder_optimizer, criterion, max_length=MAX_LENGTH,teacher_forcing_ratio=0.5):
    encoder_hidden = encoder.initHidden()

    encoder_optimizer.zero_grad()
    decoder_optimizer.zero_grad()

    input_length = input_tensor.size(0)
    target_length = target_tensor.size(0)

    loss = 0

    for ei in range(input_length):
        encoder_output, encoder_hidden = encoder(
            input_tensor[ei], encoder_hidden)

    decoder_input = torch.tensor([[SOS_token]], device=device)

    decoder_hidden = encoder_hidden

    use_teacher_forcing = True if random.random() < teacher_forcing_ratio else False

    for di in range(target_length):
        decoder_output, decoder_hidden = decoder(
            decoder_input, decoder_hidden)
        if use_teacher_forcing: # Task: Teacher forcing: Feed the target as the next decoder input

        else:                   # Task: Without teacher forcing: use its own predictions as the next input

        loss += criterion(decoder_output, target_tensor[di])
        # Task: Stop if terminate token (EOS) is generated


    # Task: Perform gradient step


    return loss.item() / target_length
```

```python
import time
import math


def asMinutes(s):
    m = math.floor(s / 60)
    s -= m * 60
    return '%dm %ds' % (m, s)


def timeSince(since, percent):
    now = time.time()
    s = now - since
    es = s / (percent)
    rs = es - s
    return '%s (- %s)' % (asMinutes(s), asMinutes(rs))
```

```python
def trainIters(encoder, decoder, n_iters, print_every=1000, plot_every=100, learning_rate=0.01):
    start = time.time()
    plot_losses = []
    print_loss_total = 0  # Reset every print_every
    plot_loss_total = 0  # Reset every plot_every

    encoder_optimizer = optim.SGD(encoder.parameters(), lr=learning_rate)
    decoder_optimizer = optim.SGD(decoder.parameters(), lr=learning_rate)
    training_pairs = [tensorsFromPair(random.choice(pairs))
                      for i in range(n_iters)]
    criterion = nn.NLLLoss()

    for iter in range(1, n_iters + 1):
        training_pair = training_pairs[iter - 1]
        input_tensor = training_pair[0]
        target_tensor = training_pair[1]

        loss = train(input_tensor, target_tensor, encoder,
                     decoder, encoder_optimizer, decoder_optimizer, criterion)
        print_loss_total += loss
        plot_loss_total += loss

        if iter % print_every == 0:
            print_loss_avg = print_loss_total / print_every
            print_loss_total = 0
            print('%s (%d %d%%) %.4f' % (timeSince(start, iter / n_iters),
                                         iter, iter / n_iters * 100, print_loss_avg))

        if iter % plot_every == 0:
            plot_loss_avg = plot_loss_total / plot_every
            plot_losses.append(plot_loss_avg)
            plot_loss_total = 0

    showPlot(plot_losses)

```

```python
import matplotlib.pyplot as plt
plt.switch_backend('agg')
import matplotlib.ticker as ticker
import numpy as np


def showPlot(points):
    plt.figure()
    fig, ax = plt.subplots()
    # this locator puts ticks at regular intervals
    loc = ticker.MultipleLocator(base=0.2)
    ax.yaxis.set_major_locator(loc)
    plt.plot(points)
```

<!-- grade_id: ce4830, points: 0, grade: false, solution: true -->
```python
def evaluate(encoder, decoder, sentence, max_length=MAX_LENGTH):
    with torch.no_grad():
        input_tensor = tensorFromSentence(input_lang, sentence)
        input_length = input_tensor.size()[0]
        encoder_hidden = encoder.initHidden()


        decoded_words = []
        # Translate the sentence. Substitute EOS_token with "<EOS>" and append to the end of decoded_words

        return decoded_words
```

```python
def evaluateRandomly(encoder, decoder, n=10):
    for i in range(n):
        pair = random.choice(pairs)
        print('>', pair[0])
        print('=', pair[1])
        output_words = evaluate(encoder, decoder, pair[0])
        output_sentence = ' '.join(output_words)
        print('<', output_sentence)
        print('')
```

```python
hidden_size = 256
encoder = EncoderRNN(input_lang.n_words, hidden_size).to(device)
decoder = DecoderRNN(hidden_size, output_lang.n_words).to(device)

trainIters(encoder, decoder, 75000, print_every=1000)
```

Output:
```
0m 58s (- 13m 37s) (5000 6%) 2.9403
1m 51s (- 12m 4s) (10000 13%) 2.3534
2m 45s (- 11m 2s) (15000 20%) 2.0395
3m 39s (- 10m 2s) (20000 26%) 1.7804
4m 33s (- 9m 6s) (25000 33%) 1.5749
5m 26s (- 8m 10s) (30000 40%) 1.3736
6m 20s (- 7m 15s) (35000 46%) 1.2614
7m 14s (- 6m 20s) (40000 53%) 1.1052
8m 8s (- 5m 25s) (45000 60%) 0.9860
9m 3s (- 4m 31s) (50000 66%) 0.8957
9m 57s (- 3m 37s) (55000 73%) 0.7844
10m 51s (- 2m 42s) (60000 80%) 0.6907
11m 45s (- 1m 48s) (65000 86%) 0.6135
12m 40s (- 0m 54s) (70000 93%) 0.5846
13m 34s (- 0m 0s) (75000 100%) 0.5152
```

```python
evaluateRandomly(encoder, decoder)
```

Output:
```
> tu es gentil .
= you re kind .
< you re a . <EOS>

> il est endormi .
= he s asleep .
< he s asleep . <EOS>

> vous etes cruel .
= you re cruel .
< you re cruel . <EOS>

> nous sommes des prisonnieres .
= we re prisoners .
< we re prisoners . <EOS>

> je suis desolee de vous ennuyer si souvent .
= i m sorry to bother you so often .
< i m sorry to bother you so often . <EOS>

> je suis heureux que tu puisses venir .
= i m glad that you can come .
< i m glad that you can . <EOS>

> elle t attend chez elle .
= she s waiting for you at home .
< she s waiting for you at home . <EOS>

> nous sommes pieges !
= we re trapped !
< we re trapped ! <EOS>

> elle ecrit une lettre maintenant .
= she is writing a letter now .
< she is writing a letter now . <EOS>

> nous sommes toutes differentes .
= we re all different .
< we re all different . <EOS>
```

```python

```